# Homework 3: Table Manipulation and Visualization

**Reading**: Textbook chapter [7](https://umass-data-science.github.io/190fwebsite/textbook/07/visualization/).

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests. Each time you start your server, you will need to execute this cell again to load the tests.

Homework 3 is due Tuesday, 9/27 at 8:00am. Start early so that you can come to office hours if you're stuck. Check the website for the office hours schedule. Late work will not be accepted as per the [policies](https://umass-data-science.github.io/190fwebsite/policy/) of this course. 

Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells. 

```python
import otter
import numpy as np
from datascience import *
grader = otter.Notebook()

# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## 1. Differences between Universities, Part II


**Question 1.1** Suppose you're choosing a university to attend, and you'd like to *quantify* how *dissimilar* any two universities are.  You rate each university you're considering on several numerical traits.  You decide on a very detailed list of 1000 traits, and you measure all of them!  Some examples:

* The cost to attend (per year)
* The average Yelp review of nearby Thai restaurants
* The USA Today ranking of the Medical school
* The USA Today ranking of the Engineering school

You decide that the dissimilarity between two universities is the *total* of the differences in their traits.  That is, the dissimilarity is:

* the **sum** of
* the absolute values of
* the 1000 differences in their trait values.

In the next cell, we've loaded arrays containing the 1000 trait values for UMass and Harvard.  Compute the dissimilarity (according to the above method) between UMass and Harvard.  Call your answer `dissimilarity`.  Use a single line of code to compute the answer.

*Note:* The data we're using aren't real -- we made them up for this exercise, except for the cost-of-attendance numbers. 

```python
UMass = Table.read_table("UMass.csv").column("Trait value")
Harvard = Table.read_table("Harvard.csv").column("Trait value")

dissimilarity = ...
dissimilarity
```

```python
grader.check("q1.1")
```

**Question 1.2** Why do we sum up the absolute values of the differences in trait values, rather than just summing up the differences?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

##### Weighing the traits
After computing dissimilarities between several schools, you notice a problem with your method: the scale of the traits matters a lot.

Since schools cost tens of thousands of dollars to attend, the cost-to-attend trait is always a much bigger *number* than most other traits.  That makes it affect the dissimilarity a lot more than other traits.  Two schools that differ in cost-to-attend by \$900, but are otherwise identical, get a dissimilarity of 900.  But two schools that differ in graduation rate by .9 (a huge difference!), but are otherwise identical, get a dissimilarity of only .9.

One way to fix this problem is to assign different "weights" to different traits.  For example, we could fix the problem above by multiplying the difference in the cost-to-attend traits by .001, so that a difference of 900 in cost-to-attend results in a dissimilarity of $900 \times .001$, or $.9$.

Here's a revised method that does that for every trait:

1. For each trait, subtract the two schools' trait values.
2. Then take the absolute value of that difference.
3. *Now multiply that absolute value by a trait-specific number, like .001 or 2.*
4. Now sum the 1000 resulting numbers.

**Question 1.3** Suppose you've already decided on a weight for each trait.  These are loaded into an array called `weights` in the cell below.  `weights.item(0)` is the weight for the first trait, `weights.item(1)` is the weight for the second trait, and so on.  Use the revised method to compute a revised dissimilarity between UMass and Harvard.

*Hint:* Using array arithmetic, your answer should be almost as short as in question 1.

```python
weights = Table.read_table("weights.csv").column("Weight")

revised_dissimilarity = ...
revised_dissimilarity
```

```python
grader.check("q1.3")
```

## 2. Unemployment


The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  Below we've loaded data on unemployment in the United States.  There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  (Economists are great at marketing.)

The source of the data is [here](https://fred.stlouisfed.org/categories/33509).

**Question 2.1** The data are in a CSV file called `unemployment.csv`.  Load that file into a table called `unemployment`.

```python
unemployment = ...
unemployment
```

```python
grader.check("q2.1")
```

**Question 2.2** Sort the data in decreasing order by NEI, naming the sorted table `by_nei`.  Create another table called `by_nei_pter` that's sorted in decreasing order by NEI-PTER instead.

```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("q2.2")
```

**Question 2.3** Use `.take` to make a table containing the data for the 10 quarters when NEI was greatest.  Call that table `greatest_nei`.

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("q2.3")
```

**Question 2.4** It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI-PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  Compute an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

*Note:* Use the original `unemployment` table for this.

```python
pter = ...
pter
```

```python
grader.check("q2.4")
```

**Question 2.5** Add `pter` as a column to `unemployment` (named "PTER") and sort the resulting table by that column in decreasing order.  Call the table `by_pter`.

Try to do this with a single line of code, if you can.

```python
by_pter = ...
by_pter
```

```python
grader.check("q2.5")
```

**Question 2.6** Create a line plot of the PTER over time. To do this, first add the `year` array and the `pter` array to the `unemployment` table; label these columns "Year" and "PTER", respectively. Then, generate a line plot using one of the table methods you've learned in class.

_Hint_: Remember that the unemployment table contains data collected quaterly (four times per year). So the year column should contain quarters of a year to plot the information correctly: 1994, 1994.25, 1994.5, 1994.75, 1995, ....

```python
year = ...
unemployment = ...
unemployment.plot("Year", "PTER")
```

```python
grader.check("q2.6")
```

**Question 2.7** Were PTER rates high during or directly after the Great Recession (that is to say, were PTER rates particularly high in the years 2008 through 2011 or so)? Assign highPTER to `True` if you think PTER rates were high in this period, and `False` if you think they weren't. 

```python
highPTER = ...
```

```python
grader.check("q2.7")
```

## 3. Birth Rates


The following table gives census-based population estimates for each state on July 1, 2015 and July 1, 2016. The last four columns describe the components of the estimated change in population during this time interval. *For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.*

```python
# Don't change this cell; just run it.
# From http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv
# See http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf
#     for column descriptions. (As of Feb 2017, no descriptions were posted for 2010-2016.)
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled(2, '2015').relabeled(3, '2016')
pop = pop.relabeled(4, 'BIRTHS').relabeled(5, 'DEATHS')
pop = pop.relabeled(6, 'MIGRATION').relabeled(7, 'OTHER')
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

**Question 3.1** Assign `us_birth_rate` to the total US annual birth rate during this time interval. The annual birth rate for a year-long period is the number of births in that period as a proportion of the population at the start of the period.

```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("q3.1")
```

**Question 3.2** Assign `fastest_growth` to an array of the names of the five states with the fastest population growth rates in *descending order of growth rate*. First, create a new version of the `pop` table, called `growth_rates`, which includes a column with the growth rate of each state. Making intermediate tables can improve the readability of the code and make it easier to follow when revisting at a later time.

```python
growth_rates = ...
fastest_growth = ...
fastest_growth
```

```python
grader.check("q3.2")
```

**Question 3.3** Assign `movers` to the number of states for which the absolute annual rate of migration was higher than 1%. The annual rate of migration for a year-long period is the net number of migrations (in and out) as a proportion of the population at the start of the period. The `MIGRATION` column contains estimated annual net migration counts by state.

```python
#This may take more lines of code than just the prompt
movers = ...
movers
```

```python
grader.check("q3.3")
```

**Question 3.4** Assign `west_births` to the total number of births that occurred in region 4 (the Western US). *Hint:* Make sure you double check the type of the values in the region column.

```python
west_births = ...
west_births
```

```python
grader.check("q3.4")
```

**Question 3.5** Assign `less_than_west_births` to the number of states that had a total population in 2016 that was smaller than the *number of babies born in region 4 (the Western US)* during this time interval.

```python
less_than_west_births = ...
less_than_west_births
```

```python
grader.check("q3.5")
```

**Question 3.6**  In the code cell below, create a visualization that will show if there is an association between birth rate and death rate during this time interval. It may be helpful to create an intermediate table here.

<!-- BEGIN QUESTION -->



```python
# Generate your chart in this cell
us_birth_rate = ...
us_death_rate = ...
association_table = ...
association_table.plot("US_BIRTH_RATE","US_DEATH_RATE")
```

<!-- END QUESTION -->

**Question 3.7** `True` or `False`: There is an association between birth rate and death rate during this time interval. 

Assign your answer to `assoc` in the cell below. 

```python
assoc = ...
```

```python
grader.check("q3.7")
```

## 4. Marginal Histograms


Consider the following scatter plot: ![](scatter.png)

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

**Question 4.1** Match each of the following histograms to the code that produced them.

**Histogram A:** 
![](var1.png)

-------------------------------------------------------------------------- 

**Histogram B:** 
![](var2.png)

Given the scatter plot above, which histogram would this line of code create: `t.hist('x')`

Answer either 'A' or 'B' exactly

```python
xHist = ...
```

```python
grader.check("q4.1.1")
```

Given the scatter plot above, which histogram would this line of code create: `t.hist('y')`

Answer either 'A' or 'B' exactly

```python
yHist = ...
```

```python
grader.check("q4.1.2")
```

## 5. Submission


Homework 3 complete!  Be sure to run the notebook tests and verify that they all pass, **your homework will need to pass more tests in gradescope**. Choose **Save and Checkpoint** from the **File** menu, then upload the .ipynb to gradescope to submit your work.  You may submit to gradescope as many times as you want until the deadline.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.1 (points: 1)
```
>>> assert dissimilarity > 0
```
failure_message: Remember to read the Trait value column. Also remember absolute value!


Tests for q1.3 (points: 1)
```
>>> assert revised_dissimilarity > 0
```
failure_message: Remember to read the Trait value column. Also remember absolute value!


Tests for q2.1 (points: 1)
```
>>> assert unemployment != None
```
failure_message: Are you using Table.read_table()?


Tests for q2.2 (points: 1)
```
>>> assert by_nei != None
>>> assert by_nei_pter != None
```
failure_message: Use .sort()


Tests for q2.3 (points: 1)
```
>>> assert greatest_nei != None
```
failure_message: Remember to take elements 1-10 by passing an array of indecies 1-10


Tests for q2.4 (points: 1)
```
>>> assert type(pter).__module__ == np.__name__
>>> assert pter.shape[0] == 90
```
failure_message: Did you compare two columns from the unemployment table?


Tests for q2.5 (points: 1)
```
>>> by_pter_t = unemployment.with_column("PTER", pter).sort("PTER", descending = True).to_df()
>>> assert by_pter_t.shape == by_pter.to_df().shape
```
failure_message: Here you should be adding a column to unemployment by using .with_column and then using .sort on that new column


Tests for q2.6 (points: 1)
```
>>> year_t = 1994 + np.arange(by_pter.num_rows)/4
>>> pter_t = unemployment.column("NEI-PTER") - unemployment.column("NEI")
>>> assert np.array_equal(year_t, unemployment.column("Year"))
>>> assert np.array_equal(pter_t, unemployment.column("PTER"))
```
failure_message: Did you add the new columns correctly? Did you name them correctly? Did you generate the correct year increments?


Tests for q2.7 (points: 1)
```
>>> assert isinstance(highPTER, bool)
```
failure_message: This shoudl either be true or false


Tests for q3.1 (points: 1)
```
>>> assert isinstance(float(us_birth_rate), float)
>>> assert float(us_birth_rate) != (sum(pop.column("BIRTHS")) / sum(pop.column("2016")))
```
failure_message: Population at the START of the period. Also this should be a single number.


Tests for q3.2 (points: 1)
```
>>> assert fastest_growth.shape[0] == 5
```
failure_message: Do not change the pop table! Did you .take from the sorted table you created?


Tests for q3.3 (points: 1)
```
>>> #The cast to type int is incase they end up with np.int type or any other various type
>>> assert isinstance(int(movers), int)
```
failure_message: Your answer should be an integer


Tests for q3.4 (points: 1)
```
>>> #The cast to type int is incase they end up with np.int type or any other various type
>>> assert isinstance(int(west_births), int)
>>> assert 900000<int(west_births)<1000000
```
failure_message: Your answer should be an integer. Make sure you calculated the total number of births in region 4


Tests for q3.5 (points: 1)
```
>>> #The cast to type int is incase they end up with np.int type or any other various type
>>> assert isinstance(int(less_than_west_births), int)
>>> assert 0<int(less_than_west_births)<10
```
failure_message: Your answer should be an integer. Make sure you calculated the number of states that had a total population in 2016 that was less than west_births


Tests for q3.7 (points: 1)
```
>>> assert isinstance(assoc, bool)
```
failure_message: Your answer must be True or False


Tests for q4.1.1 (points: 0.5)
```
>>> assert xHist == 'A' or 'B'
```
failure_message: Your answer must be 'A' or 'B'


Tests for q4.1.2 (points: 0.5)
```
>>> #The cast to type int is incase they end up with np.int type or any other various type
>>> assert yHist == 'A' or 'B'
```
failure_message: Your answer must be 'A' or 'B'
