# Lab 5: Simulations

This lab is to be taken on **Friday 10/7 and Friday 10/14**. Students should submit the lab on Gradescope when done.

Welcome to lab 5! This week, we will go over iteration and simulations, and introduce the concept of randomness. All of this material is covered in [Chapter 9](https://umass-data-science.github.io/190fwebsite/textbook/09/randomness/) and [Chapter 10](https://umass-data-science.github.io/190fwebsite/textbook/10/sampling-and-empirical-distributions/) of the textbook. 

The data used in this lab will contain salary data and statistics for basketball players from the 2014-2015 NBA season. This data was collected from sports analytic sites [basketball-reference](http://www.basketball-reference.com) and [spotrac](http://www.spotrac.com).

First, set up the tests and imports by running the cell below.

```python
import numpy as np
from datascience import *

# These lines set up graphing capabilities.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)


import otter
grader = otter.Notebook()
```

## 1. Nachos and Conditionals

In Python, Boolean values can either be `True` or `False`. We get Boolean values when using comparison operators, among which are `<` (less than), `>` (greater than), and `==` (equal to). For a complete list, refer to [Booleans and Comparison](https://umass-data-science.github.io/190fwebsite/textbook/09/randomness/) at the start of Chapter 9.

Run the cell below to see an example of a comparison operator in action.

```python
3 > 1 + 1
```

We can even assign the result of a comparison operation to a variable.

```python
result = 10 / 2 == 5
result
```

Arrays are compatible with comparison operators. The output is an array of boolean values.

```python
make_array(1, 5, 7, 8, 3, -1) > 3
```

Waiting on the dining table just for you is a hot bowl of nachos! Let's say that whenever you take a nacho, it will have cheese, salsa, both, or neither (just a plain tortilla chip). 

Using the function call `np.random.choice(array_name)`, let's simulate taking nachos from the bowl at random. Start by running the cell below several times, and observe how the results change.

```python
nachos = make_array('cheese', 'salsa', 'both', 'neither')
np.random.choice(nachos)
```

**Question 1.** Assume we took ten nachos at random, and stored the results in an array called `ten_nachos` as done below. Find the number of nachos with only cheese using code (do not hardcode the answer).  

*Hint:* Our solution involves a comparison operator and the `np.count_nonzero` method.

```python
ten_nachos = make_array('neither', 'cheese', 'both', 'both', 'cheese', 'salsa', 'both', 'neither', 'cheese', 'both')
number_cheese = ...
number_cheese
```

```python
grader.check("q1.1")
```

**Conditional Statements**

A conditional statement is made up of many lines that allow Python to choose from different alternatives based on whether some condition is true.

Here is a basic example.

```
def sign(x):
    if x > 0:
        return 'Positive'
```

How the function works is if the input `x` is greater than `0`, we get the string `'Positive'` back.

If we want to test multiple conditions at once, we use the following general format.

```
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```

Only one of the bodies will ever be executed. Each `if` and `elif` expression is evaluated and considered in order, starting at the top. As soon as a true value is found, the corresponding body is executed, and the rest of the expression is skipped. If none of the `if` or `elif` expressions are true, then the `else body` is executed. For more examples and explanation, refer to [Section 9.1](https://umass-data-science.github.io/190fwebsite/textbook/09/1/conditional-statements/).

**Question 2.** Complete the following conditional statement so that the string `'More please'` is assigned to `say_please` if the number of nachos with cheese in `ten_nachos` is less than `5`.

*Hint*: You should not have to directly reference the variable `ten_nachos`, you can use 
`number_cheese`.

```python
say_please = '?'

if ...:
    say_please = 'More please'
    
say_please
```

```python
grader.check("q1.2")
```

**Question 3.** Write a function called `nacho_reaction` that returns a string based on the type of nacho passed in as an argument. From top to bottom, the conditions should correspond to: `'cheese'`, `'salsa'`, `'both'`, `'neither'`.  

```python
def nacho_reaction(nacho):
    if nacho == nachos[0]:
        return 'Cheesy!'
    # next condition should return 'Spicy!' for salsa
    # next condition should return 'Wow!' for both
    # next condition should return 'Meh.' for neither

spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

```python
grader.check("q1.3")
```

**Question 4.** Add a column `'Reactions'` to the table `ten_nachos_reactions` that consists of reactions for each of the nachos in `ten_nachos`. 

*Hint:* Use the `apply` method. 

```python
ten_nachos_reactions = Table().with_column('Nachos', ten_nachos)
ten_nachos_reactions= ...
ten_nachos_reactions
```

```python
grader.check("q1.4")
```

**Question 5.** Using code, find the number of `'Wow!'` reactions for the nachos in `ten_nachos_reactions`.

```python
number_wow_reactions = ...
number_wow_reactions
```

```python
grader.check("q1.5")
```

**Question 6.** Complete the function `both_or_neither`, which takes in a table of nachos with reactions (just like the one from Question 4) and returns `'Wow!'` if there are more nachos with both cheese and salsa, or `'Meh.'` if there are more nachos with neither. If there are an equal number of each, return `'Okay!'`.

```python
def both_or_neither(nacho_table):
    reactions = ...
    number_wow_reactions = ...
    number_meh_reactions = ...

    ...
        return 'Wow!'
    # next condition should return 'Meh.'
    ...
        ...
    # next condition should return 'Okay!'
    ...
        ...

many_nachos = Table().with_column('Nachos', np.random.choice(nachos, 250))
many_nachos = many_nachos.with_column('Reactions', many_nachos.apply(nacho_reaction, 'Nachos'))
result = both_or_neither(many_nachos)
result
```

```python
grader.check("q1.6")
```

## 2. Simulations
Using a `for` statement, we can perform a task multiple times. This is known as iteration. Here, we'll simulate drawing different suits from a deck of cards. 

```python
suits = make_array("♤", "♡", "♢", "♧")

draws = make_array()

repetitions = 6

for i in np.arange(repetitions):
    draws = np.append(draws, np.random.choice(suits))

draws
```

In the example above, the `for` loop appends a random draw to the `draws` array for every number in `np.arange(repetitions)`. 

"Here's a nice way to think of what we did above. We had a deck of 4 cards of different suits, we randomly drew one card, saw the suit, kept track of it in `draws`, and put the card back into the deck. We repeated this for a total of 6 times without having to repeat code, thanks to the for loop. We simulated this experiment using a `for` loop. 

Another use of iteration is to loop through a set of values. For instance, we can print out all of the colors of the rainbow.


```python
rainbow = make_array("red", "orange", "yellow", "green", "blue", "indigo", "violet")

for color in rainbow:
    print(color)
```

We can see that the indented part of the `for` loop, known as the body, is executed once for each item in `rainbow`. Note that the name `color` is arbitrary; we could easily have named it something else. The important thing is we stay consistent throughout the for loop. 

```python
for another_name in rainbow:
    print(another_name)
```

In general, however, we would like the variable name to be somewhat informative. 

**Question 1.** Clay is playing darts. His dartboard contains ten equal-sized zones with point values from 1 to 10. Write code that simulates his total score after 1000 dart tosses. Make sure to use a `for` loop.

*Hint:* There are a few steps to this problem (and most simulations): 
1. Figuring out the big picture of what we want to simulate (the total score after 1000 dart tosses)
2. Deciding the possible values you can take in the experiment (point values in this case) and simulating one example (throwing one dart)
3. Deciding how many times to run through the experiment (1000 tosses in our case) and keeping track of the total information of each time you ran through the experiment (the total score in this case)
4. Coding up the whole simulation!

```python
possible_point_values = ...
tosses = 1000
total_score = ...

# a for loop would be useful here


total_score
```

```python
grader.check("q2.1")
```

**Question 2.** In the following cell, we've loaded the text of _Pride and Prejudice_ by Jane Austen, split it into individual words, and stored these words in an array. Using a `for` loop, assign `longer_than_five` to the number of words in the novel that are more than 5 letters long.

*Hint*: You can find the number of letters in a word with the `len` function.

```python
austen_string = open('Austen_PrideAndPrejudice.txt', encoding='utf-8').read()
p_and_p_words = np.array(austen_string.split())

longer_than_five = ...

# a for loop would be useful here
...


longer_than_five
```

```python
grader.check("q2.2")
```

**Question 3.** Using simulation with 10,000 trials, assign `chance_of_all_different` to an estimate of the chance that if you pick three words from Pride and Prejudice uniformly at random (with replacement), they all have different lengths. 

*Hint*: Remember that `!=` only checks for non-equality between two items, not three. However, you can use `!=` more than once in the same line. You can also use `and` to join several conditional together.


```python
trials = 10000
different = ...

...
    ...
        
chance_of_all_different = ...

chance_of_all_different
```

```python
grader.check("q2.3")
```

**Question 4** LeBron James is drafting Basketball Players for his NBA Fantasy League. He chooses 10 times randomly from a list of players, and drafts the player regardless of whether the player has been chosen before (You could have 10 Kevin Durant's on a team!). Count how many times John Wall is chosen in a version of LeBron's draft.

This can be done with or without a loop

*Hint:* Remember `np.random.choice`

```python
players = ["John Wall", "Kevin Durant", "Kyrie Irving", "Joel Embiid", "Russell Westbrook"]
draft_picks = ...


num_wall = ...


num_wall
```

```python
grader.check("q2.4")
```

Congratulations, you're done with Lab 5!  Be sure to **submit on gradescope**

**The rest of this notebook is optional. Please make sure you submit the above sections! The questions in this section are good practice for Sampling, which is included on the next assignment.**

It is highly highly recormended you at least attempt a few of the following questions

## 3. Sampling 


The questions in this data set use NBA player statistics. Run the cell below to load the player and salary data. 

```python
player_data = Table().read_table("player_data.csv")
salary_data = Table().read_table("salary_data.csv")
full_data = salary_data.join("PlayerName", player_data, "Name")
# The show method immediately displays the contents of a table. 
# This way, we can display the top of two tables using a single cell.
player_data.show(3)
salary_data.show(3)
full_data.show(3)
```

Rather than getting data on every player, imagine that we had gotten data on only a smaller subset of the players.  For 492 players, it's not so unreasonable to expect to see all the data, but usually we aren't so lucky.  Instead, we often make *statistical inferences* about a large underlying population using a smaller sample.

To save typing and increase the clarity of your code, we will package the loading and analysis code into two functions. This will be useful in the rest of the lab as we will repeatedly need to create histograms and collect summary statistics from that data.

**Question 3.1**. Complete the `histograms` function, which takes a table with columns `Age` and `Salary` and draws a histogram for each one. Use the min and max functions to pick the bin boundaries so that all data appears for any table passed to your function. Use the same bin widths as before (1 year for `Age` and $1,000,000 for `Salary`).

*Hint*: When creating the bins for the the histograms, think critically about what the stop argument should be for `np.arange`. Histograms are inclusive on the left hand side of the interval, but not the right. So, if we have a maximum age of 80, we need a 80-81 bin in order to capture this in the histogram. 

```python
def histograms(t):
    ages = t.column('Age')
    salaries = t.column('Salary')
    age_bins = ...
    salary_bins = ...
    t.hist('Age', bins=age_bins, unit='year')
    t.hist('Salary', bins=salary_bins, unit='$')
    return age_bins # Keep this statement so that your work can be checked
    
histograms(full_data)
print('Two histograms should be displayed below')
```

```python
grader.check("q3.1")
```

**Question 3.2**. Create a function called `compute_statistics` that takes a Table containing ages and salaries and:
- Draws a histogram of ages
- Draws a histogram of salaries
- Returns a two-element array containing the average age and average salary

You can call your `histograms` function to draw the histograms!

```python
def compute_statistics(age_and_salary_data):
    
    ages = ...
    salaries = ...
    ...
    ...

full_stats = compute_statistics(full_data)
```

```python
grader.check("q3.2")
```

### Convenience sampling
One sampling methodology, which is **generally a bad idea**, is to choose players who are somehow convenient to sample.  For example, you might choose players from one team that's near your house, since it's easier to survey them.  This is called, somewhat pejoratively, *convenience sampling*.

Suppose you survey only *relatively new* players with ages less than 22.  (The more experienced players didn't bother to answer your surveys about their salaries.)

**Question 3.3**  Assign `convenience_sample_data` to a subset of `full_data` that contains only the rows for players under the age of 22.

```python
convenience_sample = full_data.where('Age',are.below(22))
convenience_sample
```

```python
grader.check("q3.3")
```

**Question 3.4** Assign `convenience_stats` to a list of the average age and average salary of your convenience sample, using the `compute_statistics` function.  Since they're computed on a sample, these are called *sample averages*. 

```python
convenience_stats = compute_statistics(convenience_sample)
convenience_stats
```

```python
grader.check("q3.4")
```

Next, we'll compare the convenience sample salaries with the full data salaries in a single histogram. To do that, we'll need to use the `bin_column` option of the `hist` method, which indicates that all columns are counts of the bins in a particular column. The following cell should not require any changes; just run it.

```python
def compare_salaries(first, second, first_title, second_title):
    """Compare the salaries in two tables."""
    max_salary = max(np.append(first.column('Salary'), second.column('Salary')))
    bins = np.arange(0, max_salary+1e6+1, 1e6)
    first_binned = first.bin('Salary', bins=bins).relabeled(1, first_title)
    second_binned = second.bin('Salary', bins=bins).relabeled(1, second_title)
    first_binned.join('bin', second_binned).hist(bin_column='bin')

compare_salaries(full_data, convenience_sample, 'All Players', 'Convenience Sample')
```

**Question 3.5** Does the convenience sample give us an accurate picture of the age and salary of the full population of NBA players in 2014-2015?  Would you expect it to, in general?  Before you move on, write a short answer in English below.  You can refer to the statistics calculated above or perform your own analysis.

*Write your answer here, replacing this text.*

### Simple random sampling
A more principled approach is to sample uniformly at random from the players.  If we ensure that each player is selected at most once, this is a *simple random sample without replacement*, sometimes abbreviated to "simple random sample" or "SRSWOR".  Imagine writing down each player's name on a card, putting the cards in an urn, and shuffling the urn.  Then, pull out cards one by one and set them aside, stopping when the specified *sample size* is reached.

We've produced two samples of the `salary_data` table in this way: `small_srswor_salary.csv` and `large_srswor_salary.csv` contain, respectively, a sample of size 44 (the same as the convenience sample) and a larger sample of size 100.  

The `load_data` function below loads a salary table and joins it with `player_data`.

```python
def load_data(salary_file):
    return player_data.join('Name', Table.read_table(salary_file), 'PlayerName')
```

**Question 3.6** Run the same analyses on the small and large samples that you previously ran on the full dataset and on the convenience sample.  Compare the accuracy of the estimates of the population statistics that we get from the convenience sample, the small simple random sample, and the large simple random sample.  (Just notice this for yourself -- the autograder will check your sample statistics but will not validate whatever you do to compare.)

```python
small_srswor_data = load_data("small_srswor_salary.csv")
small_stats = ...
large_srswor_data = load_data("large_srswor_salary.csv")
large_stats = ...
print('Full data stats:                 ', full_stats)
print('Small simple random sample stats:', small_stats)
print('Large simple random sample stats:', large_stats)
```

```python
grader.check("q3.6")
```

### Producing simple random samples
Often it's useful to take random samples even when we have a larger dataset available.  The randomized response technique was one example we saw in lecture.  Another is to help us understand how inaccurate other samples are.

Tables provide the method `sample()` for producing random samples.  Note that its default is to sample with replacement. To sample without replacement use `sample(with_replacement=False)`.

**Question 3.7** Produce a simple random sample of size 44 from `full_data`.  (You don't need to bother with a join this time -- just use `full_data.sample(...)` directly.  That will have the same result as sampling from `salary_data` and joining with `player_data`.)  Run your analysis on it again.  
- Are your results roughly similar to those in the small sample we provided you?  Run your code several times to get new samples.  
- How much does the average age change across samples? 
- What about average salary?

```python
my_small_srswor_data = ...
my_small_stats = ...
my_small_stats
```

*Write your answer here, replacing this text.*

**Question 3.8** As in the previous question, analyze several simple random samples of size 100 from `full_data`.  
- Do the histogram statistics seem to change more or less across samples of 100 than across samples of size 44?  
- Are the sample averages and histograms closer to their true values for age or for salary?  What did you expect to see?

```python
my_large_srswor_data = ...
...
```

```python
#Solution
def nacho_reaction(nacho):
    if ...:
        return 'Cheesy!'
    
    if nacho == nachos[0]: # SOLUTION NO PROMPT
        return 'Cheesy!' # SOLUTION NO PROMPT
    # SOLUTION
    elif nacho == nachos[1]:
        return 'Spicy!'
    elif nacho == nachos[2]:
        return 'Wow!'
    elif nacho == nachos[3]:
        return 'Meh.'
    # END SOLUTION
    # next condition should return 'Spicy!' for salsa
    # next condition should return 'Wow!' for both
    # next condition should return 'Meh.' for neither
    

spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.1 (points: 0)
```
>>> assert int(number_cheese) == 3
```
failure_message: You should be calling np.count_nonzero() and passing in the condition `ten_nachos == _` where the underscore represents what we're looking for


Tests for q1.2 (points: 0)
```
>>> assert say_please == 'More please'
```
failure_message: How many nachos with cheese are there?


Tests for q1.3 (points: 0)
```
>>> assert nacho_reaction('both') == 'Wow!'
```
failure_message: Are your conditions correct


Tests for q1.4 (points: 0)
```
>>> assert np.count_nonzero(ten_nachos_reactions.column('Reactions') == make_array('Meh.', 'Cheesy!', 'Wow!', 'Wow!', 'Cheesy!', 'Spicy!', 'Wow!', 'Meh.', 'Cheesy!', 'Wow!')) == 10
```
failure_message: Did you use the apply function properly? One or more of the reaction results could be incorrect


Tests for q1.5 (points: 0)
```
>>> assert int(number_wow_reactions) == number_wow_reactions
```
failure_message: Incorrect value for number_wow_reactions


Tests for q1.6 (points: 0)
```
>>> seven_nachos = make_array('neither', 'cheese', 'both', 'both', 'neither', 'both', 'neither')
>>> seven_nachos_reactions = Table().with_column('Nachos', seven_nachos)
>>> seven_nachos_reactions = seven_nachos_reactions.with_column('Reactions', seven_nachos_reactions.apply(nacho_reaction, 'Nachos'))
>>> assert str(both_or_neither(seven_nachos_reactions)) == 'Okay!'
>>> 
>>> six_nachos = make_array('neither', 'cheese', 'both', 'both', 'both', 'neither')
>>> six_nachos_reactions = Table().with_column('Nachos', six_nachos)
>>> six_nachos_reactions = six_nachos_reactions.with_column('Reactions', six_nachos_reactions.apply(nacho_reaction, 'Nachos'))
>>> assert str(both_or_neither(six_nachos_reactions)) == 'Wow!'
```
failure_message: Your both or neither function did not behave as expected


Tests for q2.1 (points: 0)
```
>>> assert 1000 <= int(total_score) <= 10000
```
failure_message: You should be simulating 1000 tosses that score between 1 and 10


Tests for q2.2 (points: 0)
```
>>> assert longer_than_five == 35453
```
failure_message: Did you count every word longer than 5 characters


Tests for q2.3 (points: 0)
```
>>> assert 0.58 <= float(chance_of_all_different) <= 0.68
```
failure_message: It seems your simulation is returning a value either too low or too high. Are you checking that each of the 3 words has a different length


Tests for q2.4 (points: 0)
```
>>> assert 0 <= float(num_wall) <= 10
```


Tests for q3.1 (points: 0)
```
>>> old = full_data.with_column('Age', full_data.column('Age')*3)
>>> assert len(histograms(full_data)) < len(histograms(old))
```


Tests for q3.2 (points: 0)
```
>>> assert round(float(compute_statistics(full_data)[0]), 2) == 26.54
>>> assert round(float(compute_statistics(full_data)[1]), 2) == 4269775.77
```


Tests for q3.3 (points: 0)
```
>>> assert convenience_sample.num_rows == 44
>>> assert convenience_sample.num_columns == 11
```


Tests for q3.4 (points: 0)
```
>>> assert len(convenience_stats) == 2
>>> assert round(float(convenience_stats[0]), 2) == 20.36
>>> assert round(float(convenience_stats[1]), 2) == 2383533.82
```


Tests for q3.6 (points: 0)
```
>>> assert round(float(small_stats[0]), 2) == 26.32
>>> assert round(float(large_stats[0]), 2) == 26.42
```
