# Homework 4: Functions, Histograms, and Groups

**Reading**: Textbook chapter [7.2](https://umass-data-science.github.io/190fwebsite/textbook/07/2/visualizing-numerical-distributions/) and chapter [8](https://umass-data-science.github.io/190fwebsite/textbook/08/functions-and-tables/).

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests. Each time you start your server, you will need to execute this cell again to load the tests.

Homework 4 is due **Tueday, 10/4 at 2:00pm**. Start early so that you can come to office hours if you're stuck. Check the website for the office hours schedule.
Late work will not be accepted as per the [policies](https://umass-data-science.github.io/190fwebsite/policy/) of this course. 

Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells. 

```python
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## 1. Working with Text using Functions


The following table contains the words from four chapters of Charles Dickens' [*A Tale of Two Cities*](http://www.gutenberg.org/cache/epub/98/pg98.txt).  We're going to compute some simple facts about each chapter.  Since we're performing the same computation on each chapter, it's best to encapsulate each computational procedure in a function, and then call the function several times. Run the cell to get a table with one column.

```python
# Just run this cell to load the data.
tale_chapters = Table.read_table("tale.csv")
tale_chapters
```

**Question 1.1** Write a function called `word_count` that takes a single argument, the text of a single chapter, and returns the number of words in that chapter.  Assume that words are separated from each other by spaces. 

*Hint:* Try the string method [`split`](https://docs.python.org/3/library/stdtypes.html#str.split) and the function [`len`](https://docs.python.org/3/library/functions.html#len).

```python
def word_count(text):
        
    return
    
word_count(tale_chapters.column("Chapter text").item(0))
```

```python
grader.check("q1.1")
```

**Question 1.2** Use the `apply` function on the table from before to create an array called `chapter_lengths` which contains the length of each chapter in `tale_chapters`.

```python
chapter_lengths = ...
chapter_lengths
```

```python
grader.check("q1.2")
```

**Question 1.3** Write a function called `character_count`.  It should take a string as its argument and return the number of characters in that string that aren't spaces (" "), periods ("."), exclamation marks ("!"), or question marks ("?"). Remember that `tale_chapters` is a table, and that the function takes in only the text of one chapter as input.

*Hint:* Try using the string method `replace` several times to remove the characters we don't want to count.

```python
def character_count(text):
        
    return
    
chapter_chars = tale_chapters.apply(character_count,"Chapter text")
chapter_chars
```

```python
grader.check("q1.3")
```

**Question 1.4** Write a function called `chapter_number`.  It should take a single argument, the text of a chapter from our dataset, and return the number of that chapter, as a Roman numeral.  (For example, it should return the string "I" for the first chapter and "II" for the second.)  If the argument doesn't have a chapter number in the same place as the chapters in our dataset, `chapter_number` can return whatever you like.

To help you with this, we've included a function called `text_before`.  Its documentation describes what it does.

```python
def text_before(full_text, pattern):
    """Finds all the text that occurs in full_text before the specified pattern.

    Parameters
    ----------
    full_text : str
        The text we want to search within.
    param2 : str
        The thing we want to search for.

    Returns
    -------
    str
        All the text that occurs in full_text before pattern.  If pattern
        doesn't appear anywhere, all of full_text is returned.
    
    Examples
    --------
    
    >>> text_before("The rain in Spain falls mainly on the plain.", "Spain")
    'The rain in '
    >>> text_before("The rain in Spain falls mainly on the plain.", "ain")
    'The r'
    >>> text_before("The rain in Spain falls mainly on the plain.", "Portugal")
    'The rain in Spain falls mainly on the plain.'
    """
    return np.array(full_text.split(pattern)).item(0)

def chapter_number(chapter_text):
    
    return
    
    
    
tale_chapters.apply(chapter_number,"Chapter text")
```

```python
grader.check("q1.4")
```

## 2. Uber


Below we load tables containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the [Uber Movement](https://movement.uber.com) project. The `sourceid` and `dstid` columns contain codes corresponding to start and end locations of each ride. The `hod` column contains codes corresponding to the hour of the day the ride took place. The `ride time` table contains the length of the ride, in minutes.

```python
boston = Table.read_table("boston.csv")
manila = Table.read_table("manila.csv")
print("Boston")
boston.show(4)
print("Manila")
manila.show(4)
```

**Question 2.1**
Produce histograms of all ride times in Boston and in Manila, using the given bins. Please put the code for both of them in the following cell, and put the ride times for Boston first.  

```python
bins = np.arange(0, 120, 5)

boston.hist("ride time", bins=bins)
manila.hist("ride time", bins=bins)
```

**Question 2.2** Set the two variables below to estimates of what percentage of rides are less than 10 minutes in Boston and Manila. Find your estimates by visually assessing the histograms. Your solution should consist of only mathematical operations and numbers.

<!-- BEGIN QUESTION -->



```python
boston_under_10 = ...
manila_under_10 = ...
print("Boston: ", boston_under_10, " Manila: ", manila_under_10)
```

<!-- END QUESTION -->

**Question 2.3** Comment on the main difference between the two histograms. What might be causing this?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Optional Question 2.4** The `hod` column in each table represents the hour of the day during which the Uber was called. 0 corresponds to 12-1 AM, 1 to 1-2 AM, 13 to 1-2 PM, etc. Write a function which takes in a table like `boston` or `manila`, and an `hod` number between 0 and 23, and displays a histogram of ride lengths from that hour in that city. Use the same bins as before.

<!-- BEGIN QUESTION -->



```python
def hist_for_time(tbl, hod):
    bins = np.arange(0, 120, 5)
    
    

#DO NOT DELETE THIS LINE! 
hist_for_time(boston, 12)
```

<!-- END QUESTION -->

**Optional Question 2.5** Between the two cities, which has a larger difference between Uber calls at 10 AM and 10 PM? Use the function you just created to answer this question. Include the code you need to come up with an answer below.

Assign `larger_diff` to the number 1 if the answer is Manila, and 2 if the answer is Boston. 

<!-- BEGIN QUESTION -->



```python




larger_diff = ...
```

<!-- END QUESTION -->

## 3. Faculty salaries


In the next cell, we load a dataset compiled by the [Daily Cal](http://projects.dailycal.org/paychecker/) about UC Berkeley (where this course was created). It includes information on faculty, their departments, their positions, and their gross salaries in 2015. 

```python
profs = Table.read_table("faculty.csv").where("year", are.equal_to(2015))\
                                       .drop("year", "title")\
                                       .relabeled("title_category", "position")
profs
```

We want to use this table to generate arrays with the names of each professor in each department.

**Question 3.1** Set `prof_names` to a table with two columns. The first column should be called "department" and have the name of every department once, and the second column should be called "faculty" and contain an *array* of the names of all faculty members in that department. 

*Hint:* Think about how ```group``` works: it collects values into an array and then applies a function to that array. We have defined two functions below for you, and you will need to use one of them in your call to ```group```. If you're not sure which to use, try experimenting with both!

```python
# Pick between the two functions defined below 
def identity(array):
    return np.unique(array) 

def first(array):
    return array.item(0)
```

```python
prof_names=...
prof_names
```

```python
grader.check("q3.1")
```

**Question 3.2** At the moment, the `name` column of `profs` is sorted by last name. Would the arrays you generated in the previous part be the same as if we had sorted by first name instead? When we say that two arrays are the same, we mean that they contain the same number of elements and the elements in the corresponding positions are identical. Explain your answer. If you feel you need to make certain assumptions about the data, feel free to state them in your response. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.3** Set `biggest_range_dept` to the name of the department with the largest salary range, where range is defined as the difference between the lowest and highest salaries in the department. 

*Hint:* First you'll need to define a new function `salary_range` which takes in an array of salaries and returns the salary range of the corresponding department. Think about what functions we have discussed can combine and give you this range. Then, set `department_ranges` to a table containing the names and salary ranges of each department. 

*Hint 2:* You could use `.group("department", salary_range)` to apply the salary range function after creating a table with just the departments and gross_salary from each department

```python
# Turns an array into a the difference between highest and lowest salary
def salary_range(data):
    
    
biggest_range_dept = ...
    
biggest_range_dept
```

```python
grader.check("q3.3")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.1 (points: 1)
```
>>> assert int(word_count("abc abc abc")) == 3
```
failure_message: Your function isn't returning the correct number of words. Try testing your solution using small strings and printing the results


Tests for q1.2 (points: 1)
```
>>> assert chapter_lengths.shape[0] == 4
```
failure_message: Your result should be an array of values. Use the .apply()


Tests for q1.3 (points: 1)
```
>>> assert chapter_chars.shape[0] == 4
```
failure_message: Your result should be an array of values.


Tests for q1.4 (points: 1)
```
>>> assert tale_chapters.shape[0] == 4
```
failure_message: Your result should be an array of values.


Tests for q3.1 (points: 1)
```
>>> assert prof_names["department"].shape[0] == 71
>>> assert prof_names["faculty"].shape[0] == 71
```
failure_message: Are you creating a table with the names of faculty in each department? Make sure to relable the second column to faculty


Tests for q3.3 (points: 1)
```
>>> assert type(biggest_range_dept) == str
>>> assert len(str(biggest_range_dept)) > 4
```
failure_message: biggest_range_department should just be the name of a single department
