## Examining the Therapeutic Touch

This lab is to be taken on Friday 10/7 and Friday 10/28. Students should submit the lab on Gradescope when done.

Welcome to Lab 6! After such an extensive introduction to programming for data science, we are finally moving into the section of the course where we can apply our new skils to answer real questions.  

In this lab, we'll use testing techniques that were introduced in lecture to test the idea of the therapeutic touch, the idea that some practictioners can feel and massage your human energy field. 

```python
# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets
```

### What is the Therapeutic Touch

The Therapeutic Touch (TT) is the idea that everyone can field the Human Energy Field (HEF) around individuals. Certain practictioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals. Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy". 

TT was a popular technique used throughout the 20th century that was recommended as a great way to bring balance to a person's health. 

### Emily Rosa

Emily Rosa was a 4th grade student who had wide exposure to the world of TT due to her parents. Her parents were both medical practitioners and skeptics of the idea of TT. 

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly had any interaction with a person's HEF. 

**Question 1:** Discuss with your team mate(s) how you would set up an experiment to test this.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Emily's Experiment

Emily's experiement was clean, simple, elegant, and effective. Due to her parents occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner, ask them to extend their hands through a screen (through which they can't see). On the other side would be Emily, who would flip a coin and put out either her left hand or her right hand, depending on the coin. The TT practitioner would then have to correctly answer which hand Emily put out. Overall, through 210 samples, the practitioner picked the correct hand 44% of the time. 

Emily's main goal here was to test whether or not the TT practicioners guesses were random, like a flip of the coin. 

**Question 2**: What are the null and alternative hypothesis for Emily's experiment? Discuss with students around you to come to a conclusion. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3:** Remember that the pracitioner got the correct answer 44% of the time. Under the null hypothesis, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is between 0 and 1. 

```python
expected_correct = ...
expected_correct
```

```python
grader.check("q1.3")
```

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

**Question 4:** What is a valid test statistic we can use to test our model? Assign `valid_ts` to an array of integers representing the following options: 

1. The product of the expected percent correct and the actual percent correct
2. The absolute difference of the expected percent correct and the actual percent correct
3. The sum of the expected percent correct and the actual percent correct

There may be more than one correct answer. 

```python
# Use make_array()
valid_ts = ...
valid_ts
```

```python
grader.check("q1.4")
```

**Question 5:** Define the function `test_statistic` which takes in an expected proportion and an actual proportion, and returns the value of the test statistic chosen above. Assume that you are taking in proportions, but return your answer as a percentage. 

```python
def test_statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q1.5")
```

**Question 6:** Use your newly defined function to calculate the observed test statistic from Emily's experiment where the practicioner was correct 44% of the time

```python
observed_test_statistic = ...
observed_test_statistic
```

```python
grader.check("q1.6")
```

Is this a test statistic that is likely if the null hypothesis was true? Or is the deviation from the expected proportion due to something other than chance?

In order to answer this question, we must simulate the experiment as though the null hypothesis was true, and calculate the test statistic per each simulation. 

**Question 7:** To begin simulating, we should start by creating an array which has two items in it. The first item should be the proportion of times, assuming the null model is true, a TT practictioner picks the correct hand. The second item should be the proportion of times, under the same assumption, that the TT practicioner picks the incorrect hand. Assign `model_proportions` to this array. After this, simulate, using the `sample_proportions` function, Emily running through this experiment 210 times (as done in real life), and assign the proportion of correct answers to `simulation_proportion`. Lastly, define `one_test_statistic` to the test statistic of this one simulation. 

```python
model_proportions = ...

simulation_proportion = ...
one_test_statistic = ...
one_test_statistic
```

```python
grader.check("q1.7")
```

**Question 8:** Let's now see what the distribution of test statistics is actually like under our fully specified model. Assign `simulated_test_statistics` to an array of 1000 test statistics that you simulated assuming the null hypothesis is true. 

*Hint:* This should follow the same pattern as normal simulations, in combination with the code you did in the previous problem.  

```python
num_repetitions = 1000
num_guesses = 210

simulated_test_statistics = ...


for _ in range(num_repetitions):
    ...
    
simulated_test_statistics = ...
```

```python
grader.check("q1.8")
```

Let's view the distribution of the simulated test statistics under the null, and visually compare how the observed test statistic lies against the rest. 

```python
t = Table().with_column('Simulated Test Statistics', simulated_test_statistics)
t.hist()
```

We can make a visual argument as to whether or not we believe the observed test statistic is likely to occur under the null, or we can use the definition of p-values to help us make a more formal argument. 

**Question 9:** Assume `p_value` to the integer corresponding to the correct definition of what a p-value really is. 

1. The chance, under the null hypothesis, that the test statistic is equal to the value that was observed
2. The chance, under the null hypothesis, that the test statistic is equal to the value that was observed or is even further in the direction of the alternative
3. The chance, under the alternative hypothesis, that the test statistic is equal to the value that was observed or is even further in the direction of the null 
4. The number of times, under the null hypothesis, that the test statistic is equal to the value that was observed or is even further in the direction of the alternative

```python
p_value = ...
p_value
```

```python
grader.check("q1.9")
```

**Question 10:** Using the definition above, calculate the p-value of Emily's observed value in this experiment. 

*Hint:* If our test statistic is further in the direction of the alternative, will larger value or a smaller value? 

```python
emily_p_val = ...
emily_p_val
```

```python
grader.check("q1.10")
```

If our p-value is less than or equal to .05, then this is in favor of our alternative and we reject the null hypothesis. Otherwise, we do not have enough evidence against our null hypothesis. Note that this does **not** say we side in favor with the null hypothesis and accept it, but rather, that we just fail to reject it. 

This should help you make your own conclusions about Emily Rosa's experiment. 

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medicine journals that every professional strives for. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!

[Checkout this wikipedia page to learn more about her](https://en.wikipedia.org/wiki/Emily_Rosa)

Think to yourself and be prepared to talk with your labmates and TA about the following questions: 

1. Do we reject the null hypothesis, or fail to reject it? 
2. What does this mean in terms of Emily's experiment? Do the TT practitioners answers follow an even chance model or is there something else at play? 


```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.10 (points: 0)
```
>>> assert float(np.round(emily_p_val,2)) == .08
```
failure_message: Think about how a p_value is useful to us


Tests for q1.3 (points: 0)
```
>>> assert float(expected_correct) == 0.5
```
failure_message: This should be an float between 0 and 1. Think about the probability when you flip a coin


Tests for q1.4 (points: 0)
```
>>> assert len(valid_ts) == 1 and valid_ts[0] == 2
```
failure_message: This should be an float between 0 and 1. Think about the probability when you flip a coin


Tests for q1.5 (points: 0)
```
>>> assert int(round(test_statistic(.5,.5) + test_statistic(.1,.4),1)) == 30
```
failure_message: You're taking in two decimals and returning the absolute difference between them as a percentage


Tests for q1.6 (points: 0)
```
>>> assert float(observed_test_statistic) == 6.0
```


Tests for q1.7 (points: 0)
```
>>> assert len(model_proportions) == 2
>>> assert float(model_proportions[0]) == 0.5
>>> assert 0.25<=simulation_proportion<=0.75
>>> assert 0 <= one_test_statistic <= 20
```
failure_message: Either you did something wrong, or you got very very unlucky. Try running this again


Tests for q1.8 (points: 0)
```
>>> assert len(simulated_test_statistics) == 1000
```


Tests for q1.9 (points: 0)
```
>>> assert int(p_value) == 2
```
failure_message: Think about how a p_value is useful to us
