# Homework 5: Probability, Simulation, Estimation, and Empirical Distribution

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests. Each time you start your server, you will need to execute this cell again to load the tests.

Homework 5 is due **Tuesday, 10/18 at 8:00am**. Start early so that you can come to office hours if you're stuck. Check the website for the office hours schedule. Late assignments will be accepted until **Thursday, 10/20 at 8am but you will lose up-to 5 points** as a penalty for late submission. 

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the policies page to learn more about how to learn cooperatively.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

*The tests in this notebook may not provide feedback on whether your answers are correct, only that they are in the required correct format.*

```python
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

import warnings
warnings.filterwarnings('ignore')

import otter
grader = otter.Notebook()
```

## 1. Probability


We will be testing some probability concepts that were introduced in lecture. For all of the following problems, we will introduce a problem statement and give you a proposed answer. Next, for each of the following questions, you must assign the provided variable to one of three integers. You are more than welcome to create more cells across this notebook to use for arithmetic operations or short programs to help answer the questions, but be sure to assign the provided variable to 1, 2, or 3 in the end. 

1. Assign the variable to 1 if you believe our proposed answer is too low.
2. Assign the variable to 2 if you believe our proposed answer is correct.
3. Assign the variable to 3 if you believe our proposed answer is too high.

**Question 1.1** You roll a 6-sided die 10 times. What is the chance of getting 10 sixes?

**Our proposed answer**: $$(\frac{1}{6})^{10}$$

Assign `ten_sixes` to either 1, 2, or 3 depending on if you think our answer is too low, correct, or too high. 

```python
ten_sixes = ...
ten_sixes
```

```python
grader.check("q1.1")
```

**Question 1.2** Take the same problem set-up as before, rolling a fair dice 10 times. What is the chance that every roll is less than or equal to 5?

**Our proposed answer**: $$1 - (\frac{1}{6})^{10}$$

*Hint: Solve for what the correct answer should be using pen and paper. Then use code to print your answer and the proposed answer*

Assign `five_or_less` to either 1, 2, or 3 depending on if you think our answer is too low, correct, or too high. 

```python
five_or_less = ...
five_or_less
```

```python
grader.check("q1.2")
```

**Question 1.3** Assume we are picking a lottery ticket. We must choose three distinct numbers from 1 to 100 and write them on a ticket. Next, someone picks three numbers one by one, each time without putting the previous number back in. We win if our numbers are all called. 

If we decide to play the game and pick our numbers as 12, 14, and 89, what is the chance that we win? 

**Our proposed answer:** $$(\frac{3}{100})^3$$

Assign `lottery` to either 1, 2, or 3 depending on if you think our answer is too low, correct, or too high. 

```python
lottery = ...
```

```python
grader.check("q1.3")
```

**Question 1.4** Assume we have two lists, list A and list B. List A contains the numbers [10,20,30], while list B contains the numbers [10,20,30,40]. We choose one number from list A randomly and one number from list B randomly. What is the chance that the number we drew from list A is larger than the number we drew from list B?

Our proposed solution: $$1/4$$

Assign `list_chances` to either 1, 2, or 3. 

```python
list_chances = ...
list_chances
```

```python
grader.check("q1.4")
```

## 2. Monkeys Typing Shakespeare
##### (...or at least the string "datascience")

A monkey is banging repeatedly on the keys of a typewriter. Each time, the monkey is equally likely to hit any of the 26 lowercase letters of the English alphabet, regardless of what it has hit before. There are no other keys on the keyboard.

**Question 2.1** Suppose the monkey hits the keyboard 11 times.  Compute the chance that the monkey types the sequence `datascience`.  (Call this `datascience_chance`.) Type in an arithmetic equation that Python can evalute.

```python
datascience_chance = ...
datascience_chance
```

```python
grader.check("q2.1")
```

**Question 2.2** Write a function called `simulate_key_strike`.  It should take **no arguments**, and it should return a random one-character string that is equally likely to be any of the 26 lower-case English letters. 

*Hint: look at the np functions on this list https://umass-data-science.github.io/190fwebsite/resources/*

```python
# We have provided the code below to compute a list called letters,
# containing all the lower-case English letters.  Print it if you
# want to verify what it contains.

import string
letters = list(string.ascii_lowercase)

def simulate_key_strike():
    """Simulates one random key strike."""
    ...

# An example call to your function:
simulate_key_strike()
```

```python
grader.check("q2.2")
```

**Question 2.3** Write a function called `simulate_several_key_strikes`.  It should take one argument: an integer specifying the number of key strikes to simulate. It should return a string containing that many characters, each one obtained from simulating a key strike by the monkey.

*Hint:* If you make a list or array of the simulated key strikes, you can convert that to a string by calling `"".join(key_strikes_array)` (if your array is called `key_strikes_array`).

```python
def simulate_several_key_strikes(num_strikes):
    # Fill in this function.  Our solution used several lines
    # of code.
    
    ...
    
# An example call to your function:
simulate_several_key_strikes(11)
```

```python
grader.check("q2.3")
```

**Question 2.4** Use `simulate_several_key_strikes` 1000 times, each time simulating the monkey striking 11 keys.  Compute the proportion of times the monkey types `"datascience"`, calling that proportion `datascience_proportion`.

```python
# Your solution may take more than one line.
...
datascience_proportion= ...
datascience_proportion
```

```python
grader.check("q2.4")
```

**Question 2.5** Check the value your simulation computed for `datascience_proportion`.  Is your simulation a good way to estimate the chance that the monkey types `"datascience"` in 11 strikes (the answer to question 1)?  Why or why not?

*Write your answer here, replacing this text.*

**Question 2.6** Compute the chance that the monkey types the letter `"e"` at least once in the 11 strikes.  Call it `e_chance`. Use algebra and type in an arithmetic equation that Python can evalute. 

```python
e_chance = ...
e_chance
```

```python
grader.check("q2.6")
```

**Question 7.** Do you think that a computer simulation is more or less effective to estimate `e_chance` compared to when we tried to estimate `datascience_chance` this way? Why or why not? (You don't need to write a simulation, but it is an interesting exercise.)

*Write your answer here, replacing this text.*

## 3. Sampling Basketball Players


This exercise uses salary data and game statistics for basketball players from the 2014-2015 NBA season. The data was collected from [Basketball-Reference](http://www.basketball-reference.com) and [Spotrac](http://www.spotrac.com).

Run the next cell to load the two datasets.

```python
player_data = Table.read_table('player_data.csv')
salary_data = Table.read_table('salary_data.csv')
player_data.show(3)
salary_data.show(3)
```

We would like to relate players' game statistics to their salaries, so we have computed a table called `full_data` that includes one row for each player who is listed in both `player_data` and `salary_data` using the table `join` method.  It includes all the columns from `player_data` and `salary_data`, except the `"PlayerName"` column.

```python
full_data = player_data.join("Name", salary_data, "PlayerName")
full_data
```

Basketball team managers would like to hire players who perform well but don't command high salaries.  From this perspective, a very crude measure of a player's *value* to their team is the number of points the player scored in a season for every **\$1000 of salary** (*Note*: the `Salary` column is in dollars, not thousands of dollars). For example, Al Horford scored 1156 points and has a salary of **\$12 million.** This is equivalent to 12,000 thousands of dollars, so his value is $\frac{1156}{12000}$.

**Question 3.1** Create a table called `full_data_with_value` that's a copy of `full_data`, with an extra column called `"Value"` containing each player's value (according to our crude measure).  Then make a histogram of players' values.  **Specify bins that make the histogram informative, and don't forget your units!** Remember that `hist()` takes in an optional third argument that allows you to specify the units!

*Hint*: Informative histograms contain a majority of the data and **exclude outliers**.
*Hint 2*: For help with `hist()`: http://data8.org/datascience/_autosummary/datascience.tables.Table.hist.html

<!-- BEGIN QUESTION -->



```python
full_data_with_value = ...
...
```

<!-- END QUESTION -->

Now suppose we weren't able to find out every player's salary (perhaps it was too costly to interview each player).  Instead, we have gathered a *simple random sample* of 100 players' salaries.  The cell below loads those data.

```python
sample_salary_data = Table.read_table("sample_salary_data.csv")
sample_salary_data.show(3)
```

**Question 3.2** Make a histogram of the values of the players in `sample_salary_data`, using the same method for measuring value we used in question 2.  **Use the same bins, too.**  

*Hint:* This will take several steps, including using a table join as shown above.

<!-- BEGIN QUESTION -->



```python
sample_data = ...

...
```

<!-- END QUESTION -->

Now let us summarize what we have seen.  To guide you, we have written most of the summary already.

**Question 3.3** Complete the statements below by filling in the **[SQUARE BRACKETS]**. 

*Hint 1:* For a refresher on distribution types, check out [Section 10.1](https://www.inferentialthinking.com/chapters/10/1/empirical-distributions.html)

*Hint 2:* The `hist()` table method ignores data points outside the range of its bins, but you may ignore this fact and calculate the areas of the bars using what you know about histograms from lecture. 

<!-- BEGIN QUESTION -->

The plot in question 3.1 displayed a(n) **[DISTRIBUTION TYPE]** distribution of the population of **[A NUMBER]** players.  The areas of the bars in the plot sum to **[A NUMBER]**.

The plot in question 3.2 displayed a(n) **[DISTRIBUTION TYPE]** distribution of the sample of **[A NUMBER]** players.  The areas of the bars in the plot sum to **[A NUMBER]**.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.4** For which range of values does the plot in question 3.2 better depict the distribution of the population's **player values**: 0 to 0.5, or above 0.5? Explain your answer. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Earthquakes


The next cell loads a table containing information about **every earthquake with a magnitude above 4.5** in 2017, compiled by the US Geological Survey. (source: https://earthquake.usgs.gov/earthquakes/search/)

```python
earthquakes = Table().read_table('earthquakes_2017.csv').select(['time', 'mag', 'place'])
earthquakes
```

There are a several earthquakes that occurred in 2017 that we're interested in, and generally, we won't have access to this large population. Instead, if we sample correctly, we can take a small subsample of earthquakes in this year to get an idea about the distribution of magnitudes throughout the year!

**Question 4.1** In the following lines of code, we take two different samples from the earthquake table, and calculate the mean of the magnitudes of these earthquakes. Are these samples representative of the population of earthquakes in the original table (that is, the should we expect the mean to be close to the population mean)? 

*Hint:* Consider the ordering of the `earthquakes` table. 

<!-- BEGIN QUESTION -->



```python
sample1 = earthquakes.sort('mag', descending = True).take(np.arange(100))
sample1_magnitude_mean = np.mean(sample1.column('mag'))
sample2 = earthquakes.take(np.arange(100))
sample2_magnitude_mean = np.mean(sample2.column('mag'))
[sample1_magnitude_mean, sample2_magnitude_mean]
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 4.2** Write code producing a sample that should represent the population of size 500 then take the mean of the magnitudes of the earthquakes in this sample. Assign these to `representative_sample` and `representative_mean` respectively. 

*Hint:* In class, what sort of samples can properly represent the population? 
*Hint 2:* This may help http://data8.org/datascience/_autosummary/datascience.tables.Table.sample.html?highlight=sample

```python
representative_sample = ...
representative_mean = ...
representative_mean
```

```python
grader.check("q4.2")
```

**Question 4.3** Suppose we want to figure out what the biggest magnitude earthquake was in 2017, but we are tasked with doing this only with a sample of 500 from the earthquakes table. 

To determine whether trying to find the biggest magnitude from a sample is a plausible idea, write code that simulates the maximum of a random sample of size 500 from the `earthquakes` table 5000 times. Assign your array of maximums to `maximums`. 

```python
maximums = ...
for i in np.arange(5000): 
    maximums = ...
maximums
```

```python
grader.check("q4.3")
```

```python
#Histogram of your maximums
Table().with_column('Largest magnitude in sample', maximums).hist('Largest magnitude in sample') 
```

**Question 4.4** We want to see if a random sample of size 500 is likely to help you determine the largest magnitude earthquake in the population. To help determine this, find the magnitude of the (actual) strongest earthquake in 2017.

```python
strongest_earthquake_magnitude = ...
strongest_earthquake_magnitude
```

```python
grader.check("q4.4")
```

**Question 4.5** Explain whether you believe you can accurately use a sample size of 500 to determine the maximum. What is a specific con of using the maximum as your estimator? Use the histogram above to help answer. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Assessing Gary's Models
#### Games with Gary

Our friend Gary comes over and asks us to play a game with him. The game works like this: 

> We will flip a fair coin 10 times, and if the number of heads is greater than or equal to 5, we win!
> 
> Otherwise, Gary wins.

We play the game once and after all 10 flips we observe 1 head. We are angry and accuse Gary of cheating! Gary is adamant, however, that the coin is fair.

Gary's model claims that there is an equal chance of getting heads or tails, but we do not believe him. We believe that the coin is clearly rigged, with heads being less likely than tails. 

**Question 5.1** Assign `coin_model_probabilities` to a two-item array containing the chance of heads as the first element and the chance of tails as the second element under Gary's model. Make sure your values are between 0 and 1. 

```python
# Heads is the element at index 0 of the array, tails is at index 1
coin_model_probabilities = ...
coin_model_probabilities
```

```python
grader.check("q5.1")
```

**Question 5.2** We believe Gary's model is incorrect. In particular, we believe there to be a smaller chance of heads. Which of the following statistics can we use during our simulation to test between his model and our alternative? Assign `statistic_choice` to the correct answer. 

1. The distance (absolute value) between the actual number of heads in 10 flips and the expected number of heads in 10 flips (5)
2. The expected number of heads in 10 flips
3. The actual number of heads we get in 10 flips



```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("q5.2")
```

**Question 5.3** Define the function `coin_simulation_and_statistic`, which, given a sample size and an array of model proportions (like the one you created in Q1), returns the number of heads in one simulation of flipping the coin under the model specified in `model_proportions`. 

*Hint:* Think about how you can use the function `sample_proportions`. See: http://data8.org/datascience/util.html?highlight=sample_proportions

```python
def coin_simulation_and_statistic(sample_size, model_proportions):
    '''
    sample_size: an int representing how many trials are run
    model_proportions: an array of two floats representing the probabilities of heads and tails.
                       heads is at index 0 tails is at index 1
    '''
    
    ...

coin_simulation_and_statistic(10, coin_model_probabilities)
```

```python
grader.check("q5.3")
```

**Question 5.4** Use your function from above to simulate the flipping of 10 coins 5000 times under the proportions that you specified in problem 1. Keep track of all of your statistics in an array `coin_statistics`. 

```python
coin_statistics = make_array()
repetitions = ...
flips = ...

...
    
coin_statistics
```

```python
grader.check("q5.4")
```

Let's take a look at the distribution of statistics, using a histogram. 

```python
#Draw a distribution of statistics 
Table().with_column('Coin Statistics', coin_statistics).hist(bins=np.arange(0,11))
```

**Question 5.5** Given your observed value, do you believe that Gary's model is reasonable, or is our alternative more likely? Explain your answer using the distribution drawn in the previous problem. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.1 (points: 1)
```
>>> assert 1<=int(ten_sixes)<=3
```
failure_message: You must enter an integer between 1-3


Tests for q1.2 (points: 1)
```
>>> assert 1<=int(five_or_less)<=3
```
failure_message: You must enter an integer between 1-3


Tests for q1.3 (points: 1)
```
>>> assert 1<=int(lottery)<=3
```
failure_message: You must enter an integer between 1-3


Tests for q1.4 (points: 1)
```
>>> assert 1<=int(list_chances)<=3
```
failure_message: You must enter an integer between 1-3


Tests for q2.1 (points: 1)
```
>>> assert 0<=float(datascience_chance)<=1
```
failure_message: Your answer should be a probability between 0-1


Tests for q2.2 (points: 1)
```
>>> assert len(simulate_key_strike()) == 1
```
failure_message: Are you generating 1 random character?
```
>>> assert 26 >= len(np.unique([simulate_key_strike() for i in range(500)])) >= 20
```
failure_message: It looks like your simulation doesn't use all the letters


Tests for q2.3 (points: 1)
```
>>> assert len(simulate_several_key_strikes(15)) == 15 and isinstance(simulate_several_key_strikes(15), str)
```
failure_message: are you generating num_strikes many characters? and returning a string?
```
>>> 
>>> assert 26 >= len(np.unique(list(simulate_several_key_strikes(500)))) >= 20
```
failure_message: It looks like your simulation doesn't use all the letters


Tests for q2.4 (points: 1)
```
>>> assert float(datascience_proportion) == 0.0
```
failure_message: There's a chance you just hit a 1 in a quadrillion chance


Tests for q2.6 (points: 1)
```
>>> assert 0<=float(e_chance)<=1
```
failure_message: This should be a probability 0-1


Tests for q4.2 (points: 1)
```
>>> assert representative_sample.num_rows == 500
```
failure_message: The sample should be of size 500.
```
>>> assert representative_sample.num_columns == 3
```
failure_message: Your sample should have the same columns as the original table
```
>>> assert float(representative_mean) < float(max(representative_sample.column('mag'))) and float(representative_mean) > float(min(representative_sample.column('mag')))
```
failure_message: The mean can't be bigger than the biggest magnitude, or smaller than the smallest


Tests for q4.3 (points: 1)
```
>>> assert len(maximums) == 5000
```
failure_message: maximums should be an array of 5000 values
```
>>> assert max(maximums) <= max(earthquakes.column('mag'))
```
failure_message: The biggest simulated maximum can't be bigger than the actual maximum!


Tests for q4.4 (points: 1)
```
>>> assert float(strongest_earthquake_magnitude) > 8.00001
```
failure_message: This should be a single float


Tests for q5.1 (points: 1)
```
>>> assert len(coin_model_probabilities) == 2 and sum(coin_model_probabilities)
```
failure_message: There should be 2 probabilities that sum to one 


Tests for q5.2 (points: 1)
```
>>> assert  1<=statistic_choice<=3
```
failure_message: This should be between 1 and 3


Tests for q5.3 (points: 1)
```
>>> for i in range(20):
...     num_heads = coin_simulation_and_statistic(10, make_array(.5,.5))
...     assert 0 <= num_heads <= 10
```
failure_message: coin_simulation_and_statistic should always return between 0-10


Tests for q5.4 (points: 1)
```
>>> assert len(coin_statistics) == 5000
>>> assert all([ 0<= k <= 10 for k in coin_statistics])
```
