# Homework 9: Linear Regression


**Reading**: Textbook chapter [15](https://umass-data-science.github.io/190fwebsite/textbook/15/prediction/).

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests. Each time you start your server, you will need to execute this cell again to load the tests.

Homework 9 is due **Monday Dec 1 at 11:30pm.** Start early so that you can come to office hours if you're stuck.

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the policies page to learn more about how to learn cooperatively.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on.

```python
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

import otter
grader = otter.Notebook()
```

## 1. Triple Jump Distances vs. Vertical Jump Heights


Does skill in one sport imply skill in a related sport?  The answer might be different for different activities.  Let us find out whether it's true for the [triple jump](https://en.wikipedia.org/wiki/Triple_jump) (an horizontal jump similar to a long jump) and the vertical jump.  Since we're learning about linear regression, we will look specifically for a *linear* association between skill in the two sports.

The following data was collected by observing 40 collegiate level soccer players.  Each athlete's distances in both jump activities were measured in centimeters. Run the cell below to load the data.

```python
# Run this cell to load the data
jumps = Table.read_table('triple_vertical.csv')
jumps
```

#### Question 1
Before running a regression, it's important to see what the data look like, because our eyes are good at picking out unusual patterns in data.  Draw a scatter plot with the triple jump distances on the horizontal axis and the vertical jump heights on vertical axis **that also shows the regression line**. 

See the documentation for [`scatter`](http://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html#datascience.tables.Table.scatter) for instructions on how to have Python draw the regression line automatically.

<!-- BEGIN QUESTION -->



```python
...
```

<!-- END QUESTION -->

#### Question 2
Does the correlation coefficient `r` look closest to 0, .5, or -.5? Explain. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Question 3
Create a function called `regression_parameters`. It takes as its argument a table with two columns.  The first column is the x-axis, and the second column is the y-axis.  It should compute the correlation between the two columns, then compute the slope and intercept of the regression line that predicts the second column from the first, in original units (centimeters).  It should return an array with three elements: the correlation coefficient of the two columns, the slope of the regression line, and the intercept of the regression line.

```python
def regression_parameters(t):
    ...
    
    # Our solution had 4 lines above this one; you may use more than that
    r = ...
    slope = ...
    intercept = ...
    return make_array(r, slope, intercept)

# When your function is finished, the next lines should
# compute the regression line predicting vertical jump 
# distances from triple jump distances. Set parameters
# to be the result of calling regression_parameters appropriately.
parameters = regression_parameters(jumps)
print('r:', parameters.item(0), '; slope:', parameters.item(1), '; intercept:', parameters.item(2))
```

```python
grader.check("q1.3")
```

#### Question 4
Let's use `parameters` to predict what certain athletes' vertical jump heights would be given their triple jump distances.

The world record for the triple jump distance is 18.29 *meters* by Johnathan Edwards. What's our prediction for what Edwards' vertical jump would be?

```python
triple_record_vert_est = ...
print("Predicted vertical jump distance: {:f} centimeters".format(triple_record_vert_est))
```

```python
grader.check("q1.4")
```

#### Question 5
Do you expect this estimate to be accurate within a few centimeters? Why or why not?

*Hint:* Compare Edwards' triple jump distance to the triple jump distances in `jumps`. Is it relatively similar to the rest of the data? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Cryptocurrencies


Imagine you're an investor in December 2017. Cryptocurrencies, online currencies backed by secure software, are becoming extremely valuable, and you want in on the action!

The two most valuable crpytocurrencies are Bitcoin (BTC) and Ethereum (ETH). Each one has a dollar price attatched to it at any given moment in time. For example, on December 1st, 2017, one BTC costs 10859.56 dollars and one ETH costs 424.64 dollars. 

**You want to predict the price of ETH at some point in time based on the price of BTC.** Below, we [load](https://www.kaggle.com/jessevent/all-crypto-currencies/data) two tables called `btc` and `eth`. Each has 5 columns:
* `date`, the date
* `open`, the value of the currency at the beginning of the day
* `close`, the value of the currency at the end of the day
* `market`, the market cap or total dollar value invested in the currency
* `day`, the number of days since the start of our data

```python
btc = Table.read_table('btc.csv')
btc
```

```python
eth = Table.read_table('eth.csv')
eth
```

#### Question 1

In the cell below, make one or two plots to investigate the opening prices of BTC and ETH as a function of time. Then comment on whether you think the values are correlated.

_Type your answer here, replacing this text._

<!-- BEGIN QUESTION -->



```python
...
```

```python
...
```

<!-- END QUESTION -->

#### Question 2

Now, calculate the correlation coefficient between the opening prices of BTC and ETH.

*Hint:* It may be helpful to define and use the function `std_units`.

```python
def std_units(arr): 
    ...

standard_btc = ...
standard_eth = ...

r = ...
r
```

```python
grader.check("q2.2")
```

#### Question 3
Regardless of your conclusions above, write a function `eth_predictor` which takes an opening BTC price and predicts the price of ETH. Again, it will be helpful to use the function `regression_parameters` that you defined earlier in this homework.

```python
def eth_predictor(btc_price):
    parameters = regression_parameters(Table().with_columns('BTC', btc[1], 'ETH', eth[1]))
    slope = parameters[1]
    intercept = parameters[2]
    return slope * btc_price + intercept
eth_predictor(10)
```

```python
grader.check("q2.3")
```

#### Question 4

Now, using the `eth_predictor` you defined in the previous question, make a scatter plot with BTC prices along the x-axis and both real and predicted ETH prices along the y-axis. The color of the dots for the real ETH prices should be different from the color for the predicted ETH prices.

Hints:
* An example of such a scatter plot is generated <a href= "https://umass-data-science.github.io/190fwebsite/textbook/15/2/regression-line/
"> here. </a>
* Think about the table that must be produced and used to generate this scatter plot. What data should the columns represent? Based on the data that you need, how many columns should be present in this table? Also, what should each row represent? Constructing the table will be the main part of this question; once you have this table, generating the scatter plot should be straightforward as usual.

<!-- BEGIN QUESTION -->



```python
predictions = eth_predictor(btc['open'])
...
```

<!-- END QUESTION -->

#### Question 5
For what values of BTC should you be more confident of your predictions? Considering the shape of the scatter plot of the true data, is the model we used reasonable? If so, what features make this model reasonable? If not, what features make it unreasonable?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Evaluating NBA Game Predictions


#### A brief introduction to sports betting

In a basketball game, each team scores some number of points.  Conventionally, the team playing at its own arena is called the "home team," and the other team is called the "away team."  The winner is the team with more points.

We can summarize what happened in a game by the "**outcome**", defined as the **the away team's score minus the home team's score**:

$$\text{outcome} = \text{points scored by the away team} - \text{points scored by the home team}$$

If this number is positive, the away team won.  If it's negative, the home team won. 

In order to facilitate betting on games, analysts at casinos try to predict the outcome of the game. This prediction of the outcome is called the **spread.**


```python
spreads = Table.read_table("spreads.csv")
spreads
```

Here's a scatter plot of the outcomes and spreads, with the spreads on the horizontal axis.

```python
spreads.scatter("Spread", "Outcome")
```

#### Question 1
Why do you think that the spread and outcome are never 0 (aside from 1 case of the spread being 0)?  

*Hint*: for those who may not be familiar with the rules of Basketball there if at the end of the game both teams are tied play continues into overtime

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Let's investigate how well the casinos are predicting game outcomes.

One question we can ask is: Is the casino's prediction correct on average? In other words, for every value of the spread, is the average outcome of games assigned that spread equal to the spread? If not, the casino would apparently be making a systematic error in its predictions.

#### Question 2
Among games with a spread between 3.5 and 6.5 (including both 3.5 and 6.5), what was the average outcome? 

*Hint:* Read the documentation for the predicate [`are.between_or_equal_to`](http://data8.org/datascience/predicates.html#datascience.predicates.are.between_or_equal_to).

```python
...
print("Average outcome for spreads around 5:", spread_5_outcome_average)
```

```python
grader.check("q3.2")
```

#### Question 3
If the average outcome for games with any given spread turned out to be exactly equal to that spread, what would the slope and intercept of the linear regression line be, in original units? Hint: If you're stuck, try drawing a picture!

```python
expected_slope_for_equal_spread = ...
expected_intercept_for_equal_spread = ...
```

```python
grader.check("q3.3")
```

#### Question 4
Implement the `standard_units` function below.  It should take an array of numbers as its argument and return an array of those numbers in standard units. 

*Hint*: You saw this earlier

```python
def standard_units(nums):
    """Return an array where every value in nums is converted to standard units."""
    ...
```

```python
grader.check("q3.4")
```

#### Question 5
Compute the correlation coefficient between outcomes and spreads using the `standard_units` function.

```python
spread_r = ...
spread_r
```

```python
grader.check("q3.5")
```

#### Question 6
Compute the slope of the least-squares linear regression line that predicts outcomes from spreads, in original units.

```python
spread_slope = ...
spread_slope
```

```python
grader.check("q3.6")
```

#### Question 7
For the "best fit" line that estimates the average outcome from the spread, the slope is less than 1. Does knowing the slope alone tell you whether the average spread was higher than the average outcome? If so, set the variable name below to `True`. If you think you need more information than just the slope of the regression line to answer that question, then respond `False`. Briefly justify your answer below. (HINT: Does the intercept matter?)

<!-- BEGIN QUESTION -->



```python
slope_implies_average_spread_above_average_outcome = ...
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.3 (points: 1)
```
>>> assert len(parameters) == 3
>>> assert np.isclose(parameters.item(0), 0.8343076972837598)
>>> assert np.isclose(parameters.item(1), 0.09295728160512184)
>>> assert np.isclose(parameters.item(2), -1.566520972963474)
```
failure_message: Your function is not returning the proper parameters. See chapter 15.2


Tests for q1.4 (points: 1)
```
>>> assert round(triple_record_vert_est,2) == round(168.452347,2)
```
failure_message: Make sure you account for 18.29 being in meters


Tests for q2.2 (points: 1)
```
>>> assert round(r,2) == round(0.9250325764148278,2)
```
failure_message: You seem to be calculating r incorrectly


Tests for q2.3 (points: 1)
```
>>> assert 1 < round(eth_predictor(10),2) < 10
```
failure_message: Your prediction seems unreasonable sized. Be sure to use regression_parameters() you defined above


Tests for q3.2 (points: 1)
```
>>> assert 4 <= spread_5_outcome_average <= 6
```
failure_message: First find every game with a spread between 3.5 and 6.5, then find the mean outcome of those games


Tests for q3.3 (points: 1)
```
>>> assert int(expected_slope_for_equal_spread) == 1
>>> assert int(expected_intercept_for_equal_spread) == 0
```
failure_message: Think about the line y=x. What is the slope? What is the intercept?


Tests for q3.4 (points: 1)
```
>>> assert np.allclose(standard_units([2, 0, 3, 6, 4]), [-0.5, -1.5,  0.,  1.5,  0.5])
```
failure_message: Look at question 2.2


Tests for q3.5 (points: 1)
```
>>> assert .25 < spread_r < 1
```
failure_message: Look at the graph and consider roughly what this value should be


Tests for q3.6 (points: 1)
```
>>> assert 0 < spread_slope < 1
```
