# Homework 2: Arrays and Tables

**Reading**: Textbook chapters [4](https://umass-data-science.github.io/190fwebsite/textbook/04/data-types/), [5](https://umass-data-science.github.io/190fwebsite/textbook/05/sequences/), and [6](https://umass-data-science.github.io/190fwebsite/textbook/06/tables/).

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided tests and necessary packages. Each time you start your server, you will need to execute this cell again.

Homework 2 is due Tuesday, 9/20 at the begining of class. Start early so that you can come to office hours if you're stuck. Check the piazza for the office hours schedule.

Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. 

```python
import otter
import numpy as np
from datascience import *
grader = otter.Notebook()
```

## 1. Creating Arrays


**Question 1.1** Make an array called `weird_numbers` containing the following numbers (in the given order):

1. -2
2. the sine of 1.2
3. 3
4. 5 to the power of the cosine of 1.2

*Hint:* `sin` and `cos` are functions in the `math` module. Remember you will need to import it then use the functions.

```python
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("q1.1")
```

**Question 1.2** Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves". Use the make_array function from class.

```python
book_title_words = ...
book_title_words
```

```python
grader.check("q1.2")
```

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string

**Question 1.3** Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this| one `without_commas`)

*Hint:* If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .

```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

```python
grader.check("q1.3")
```

## 2. Indexing Arrays


These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.  

**Question 2.1** The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q2.1")
```

**Question 2.2** The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information in the cell (the strings that are currently "???") to complete the table.

```python
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", "???", "???", "fifth"),
    "Index",                     make_array("???", "1", "2", "???", "4"),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q2.2")
```

**Question 2.3** You'll sometimes want to find the *last* element of an array.  Suppose an array has 142 elements.  What is the index of its last element? (Set index_of_last_element equal to the answer)

```python
index_of_last_element = ...
```

```python
grader.check("q2.3")
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

**Question 2.4** The cell below loads an array called `president_birth_years`.  The last element in that array is the most recent birth year of any deceased president. Assign that year to `most_recent_birth_year`.

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q2.4")
```

## 3. Basic Array Arithmetic


**Question 3.1** Multiply the integers 42, 4224, 42422424, and -250 by 157.  For this question, **don't** use arrays.

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

```python
grader.check("q3.1")
```

**Question 3.2** Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.

```python
numbers = ...
products = ...
products
```

```python
grader.check("q3.2")
```

**Question 3.3** Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the fixed products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers.

```python
fixed_products = numbers * 1577
fixed_products
```

```python
grader.check("q3.3")
```

**Question 3.4** We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a day at a climate observation station, mostly from the US.  Since they're from the US government agency NOAA, all the temperatures are in Fahrenheit.  Convert them all to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Round each result to the nearest integer using the `np.round` function.

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("q3.4")
```

**Question 3.5** The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Give your answer in Celsius!**

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("q3.5")
```

## 4. World Population


The cell below loads a table of estimates of the world population for different years, starting in 1950. The estimates come from the US Census Bureau website.

```python
world = Table.read_table("world_population.csv").select('Year', 'Population')
world.show(4)
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column(1)
population
```

In this question, you will apply some built-in Numpy functions to this array.

<img src="array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array by the element that preceeds it. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

<img src="array_cumsum.png" style="width: 700px;"/>

The cumulative sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

**Question 4.1** Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.

```python
largest_population_change = ...
largest_population_change
```

```python
grader.check("q4.1")
```

**Question 4.2** Describe in words the result of the following expression. What do the values in the resulting array represent (choose one)?

```python
np.cumsum(np.diff(population))
```

1) The total population change between consecutive years, starting at 1951.

2) The total population change between 1950 and each later year, starting at 1951.

3) The total population change between consecutive years and 1950, starting at 1951.

```python
# Assign cumulative_sum_answer to 1, 2, or 3
cumulative_sum_answer = ...
```

```python
grader.check("q4.2")
```

## 5. Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

**Question 5.1** The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q5.1")
```

**Question 5.2** Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was 74 - 62 = 12 minutes. *Hint*: You'll need an array arithmetic function [mentioned in the textbook](https://umass-data-science.github.io/190fwebsite/textbook/05/1/arrays/).

```python
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q5.2")
```

**Question 5.3** If you expected Old Faithful to erupt every hour, you would expect to wait a total of `60 * k` minutes to see `k` eruptions. Set `difference_from_expected` to an array with 272 elements, where the element at index `i` is the absolute difference between the expected and actual total amount of waiting time to see the first `i+1` eruptions.  *Hint*: You'll need to compare a cumulative sum to a range.

For example, since the first three waiting times are 79, 54, and 74, the total waiting time for 3 eruptions is 79 + 54 + 74 = 207. The expected waiting time for 3 eruptions is 60 * 3 = 180. Therefore, `difference_from_expected.item(2)` should be $|207 - 180| = 27$.

```python
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("q5.3")
```

**Question 5.4** If instead you guess that each waiting time will be the same as the previous waiting time, how many minutes would your guess differ from the actual time, averaging over every wait time except the first one.

For example, since the first three waiting times are 79, 54, and 74, the average difference between your guess and the actual time for just the second and third eruption would be $\frac{|79-54|+ |54-74|}{2} = 22.5$.

```python
average_error = ...
average_error
```

```python
grader.check("q5.4")
```

## 6. Tables


**Question 6.1** Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: "fruit name" and "count".  Give it the name `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.

```python
# Our solution uses 1 statement
fruits = ...
fruits
```

```python
grader.check("q6.1")
```

**Question 6.2** The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit.  Load it as a table named `inventory`.

```python
inventory = ...
inventory
```

```python
grader.check("q6.2")
```

**Question 6.3** Does each box at the fruit stand contain a different fruit? Set all_different to "Yes" if each box contains a different fruit or to "No" if multiple boxes contain the same fruit.

```python
all_different = ...
all_different
```

```python
grader.check("q6.3")
```

**Question 6.4** The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.

```python
sales = ...
sales
```

```python
grader.check("q6.4")
```

**Question 6.5** How many fruits did the store sell in total on that day?

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q6.5")
```

**Question 6.6** What was the store's total revenue (the total price of all fruits sold) on that day?

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q6.6")
```

**Question 6.7** Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's count, so that the "count" is the amount of fruit remaining after Saturday.

_Hint_: The ?? should be column names, and the !! should be column values either taked directly from the inventory or modified

```python
remaining_inventory = Table().with_columns(   
    "??", !!,
    "??", !!,
    "??", !! )
remaining_inventory
```

```python
grader.check("q6.7")
```

## 7. Submission


You've finished homework 2!  Be sure to run the notebook tests and verify that they all pass, **your homework will need to pass more tests in gradescope**. Choose **Save and Checkpoint** from the **File** menu, then upload the .ipynb to gradescope to submit your work.  You may submit to gradescope as many times as you want until the deadline.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1.1 (points: 1)
```
>>> assert type(weird_numbers).__module__ == np.__name__
>>> assert weird_numbers.shape[0] == 4
```
failure_message: Remember to use make_array()


Tests for q1.2 (points: 1)
```
>>> assert type(book_title_words).__module__ == np.__name__
>>> assert book_title_words.shape[0] == 3
```
failure_message: If this test failed your answer is not an array of the correct length


Tests for q1.3 (points: 1)
```
>>> assert isinstance(with_commas, str)
>>> assert isinstance(without_commas, str)
```
failure_message: If this test failed your answer is not a list of the correct length


Tests for q2.1 (points: 1)
```
>>> assert isinstance(int(third_element), int)
```
failure_message: If this test failed your answer is not the proper value from some_numbers


Tests for q2.2 (points: 1)
```
>>> df = Table().with_columns(
...     "English name for position", make_array("first", "second", "third", "fourth", "fifth"),
...     "Index",                     make_array("0", "1", "2", "3", "4"),
...     "Element",                   some_numbers).to_df()
>>> assert df.shape == elements_of_some_numbers.to_df().shape
```
failure_message: If this test failed you may have accidentaly modifed the prompt code


Tests for q2.3 (points: 1)
```
>>> index_of_last_element == int(index_of_last_element)
```
failure_message: If this test failed your answer is not a number


Tests for q2.4 (points: 1)
```
>>> assert type(president_birth_years[0]) == type(most_recent_birth_year)
```
failure_message: If this test failed your answer is not a value from the array


Tests for q3.1 (points: 1)
```
>>> assert isinstance(first_product, int)
>>> assert isinstance(second_product, int)
>>> assert isinstance(third_product, int)
>>> assert isinstance(fourth_product, int)
```
failure_message: The question asked for integers, make sure you didn't add a decimal point to any of the numbers!!


Tests for q3.2 (points: 1)
```
>>> assert type(products).__module__ == np.__name__
>>> assert products.shape[0] == 4
```
failure_message: Make sure you used make_array() and only one *


Tests for q3.3 (points: 1)
```
>>> assert type(fixed_products).__module__ == np.__name__
>>> assert fixed_products.shape[0] == 4
```
failure_message: Make sure you used make_array() and only one *


Tests for q3.4 (points: 1)
```
>>> assert type(celsius_max_temperatures).__module__ == np.__name__
```
failure_message: Soemthing went wrong, did you accidentally modify the starter code?


Tests for q3.5 (points: 1)
```
>>> no_celsius = abs((max_temperatures - 32) * (5 / 9) - min_temperatures)
>>> rounded = abs(celsius_max_temperatures - (min_temperatures - 32) * (5 / 9))
>>> assert not np.array_equal(celsius_temperature_ranges, no_celsius)
>>> assert not np.array_equal(celsius_temperature_ranges, rounded)
```
failure_message: Soemthing went wrong, did you convert to celsius before finding the difference? Make sure you didn't use the rounded temps!


Tests for q4.1 (points: 1)
```
>>> assert isinstance(int(largest_population_change), int)
```
failure_message: Did you use max()? largest_population_change should just be a single integer


Tests for q4.2 (points: 1)
```
>>> assert isinstance(int(cumulative_sum_answer), int)
>>> assert cumulative_sum_answer <=3 and cumulative_sum_answer >= 0
```
failure_message: This should be an integer between 0-3


Tests for q5.1 (points: 1)
```
>>> assert isinstance(float(shortest), float)
>>> assert isinstance(float(longest), float)
>>> assert isinstance(float(average), float)
```
failure_message: All 3 varriables should be a single number


Tests for q5.2 (points: 1)
```
>>> assert isinstance(int(biggest_decrease), int)
```
failure_message: This should be a single number


Tests for q5.3 (points: 1)
```
>>> assert difference_from_expected.shape[0] == 272
>>> assert (difference_from_expected > 0).all() == True
```
failure_message: Remember this is difference so every value should be positive use: np.abs())


Tests for q5.4 (points: 1)
```
>>> #this should fail if the answer is not a single number
>>> assert type(float(average_error)) == float
```
failure_message: This should be the mean of an array not an array


Tests for q6.1 (points: 1)
```
>>> df = Table().with_columns("fruit name", make_array("apple", "orange", "pineapple"), "count", make_array(4, 3, 3)).to_df()
>>> assert df.shape == fruits.to_df().shape
```
failure_message: Did you use Table().with_columns? Did you use make_array() for the rows? Check the inclass demos


Tests for q6.2 (points: 1)
```
>>> df = Table.read_table('inventory.csv').to_df()
>>> assert df.shape == inventory.to_df().shape
```
failure_message: Did you use TTable.read_table?


Tests for q6.3 (points: 1)
```
>>> assert all_different == "Yes" or all_different == "No"
```
failure_message: Check the case and spelling of your answers?


Tests for q6.4 (points: 1)
```
>>> df = Table.read_table('sales.csv').to_df()
>>> assert df.shape == sales.to_df().shape
```
failure_message: Did you use Table.read_table?


Tests for q6.5 (points: 1)
```
>>> assert isinstance(int(total_fruits_sold), int)
```
failure_message: Did you use Table().with_columns? Did you use make_array() for the rows? Check the inclass demos


Tests for q6.6 (points: 1)
```
>>> sol = np.sum(sales.column('count sold') * sales.column('price per fruit ($)'))
>>> assert isinstance(total_revenue, float) or type(total_revenue) == type(sol)
```
failure_message: Check your types! This should be a single value and it should have decimal places


Tests for q6.7 (points: 1)
```
>>> df = Table().with_columns(   
...     "box ID", inventory.column(0),
...     "fruit name", inventory.column(1),
...     "count", inventory.column(2)-sales.column(2)).to_df()
>>> assert df.shape == remaining_inventory.to_df().shape
>>> assert str(remaining_inventory.to_df().columns[0]) == "box ID"
```
failure_message: Check the shape of your table and the column names. Everything should be very similar to inventory.
